# Problem 1 - Introduction to river profile evolution (*3 points*)

For this exercise we will be using some Python code to plot river profiles.
The code performs many of the basic calculations needed to answer the questions below, but you will need add a few functions and make some changes to complete the exercise.

**Notice**: Closely follow the instructions! For example, you should be sure to use **exactly** the same variable names mentioned in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.

**Your score on this problem will be based on following criteria**:

- Creating a Python function for calculating drainage area
- Creating a Python function for the stream-power erosion equation
- Answering a few questions about these functions and others defined in this problem
- Including comments that explain what most lines in the code do
- Uploading your notebook to your GitHub repository for this week's exercise

## Part 1: Getting started, finding the drainage area (*1 point*)

We will be working with river profiles in this exercise, which show the elevation of a river channel along its length.
However, the stream-power erosion law relies on estimating the amount of water in the river (i.e., its discharge) using the area upstream of any given point.
Thus, your first task in this problem is to create a function that can be used to approximate the change in upstream area in a river basin as you move down stream.

We can approximate the area $A$ upstream of any point using

\begin{equation}
  \Large
  A = k_{a} x^{h_{a}}
\end{equation}

where $k_{a}$ is a distance/area scaling coefficient with units of m$^{1/3}$ and $h_{a}$ is a scaling exponent.
Typical values of $k_{a}$ and $h_{a}$ are 6.69 and 1.69, respectively.

- Run the two Python cells below to get the plotting set.
- Create a function `drainage_area` that can be used to estimate the upstream drainage area for a river profile of length $x$.
    - Your function should calculate the area (which you could call `area`) using the equation above.
    - The first value of `area` should not be 0.0, but rather you can manually set it to be half of the second value in `area`.
    - Your function should return the area array (called `area` if you follow the example name above).

```python
# Import NumPy, Matplotlib and the Matplotlib animation modules
import numpy as np
import pandas as pd
import pandas_bokeh
```

```python
# Enable interactive plotting/animations
%matplotlib inline

# Use Pandas-Bokeh for plotting
pandas_bokeh.output_notebook()
pd.set_option('plotting.backend', 'pandas_bokeh')
```

<!-- grade_id: problem_1_ag_p1_drainage_area, points: 0, grade: false, solution: true -->
```python
def drainage_area(x, ka=6.69, ha=1.69):
    """Calculate estimated drainage area upstream.
    
    Keyword arguments:
    x -- an array of distances along the river profile
    ka -- distance/area scaling coefficient (units: m**0.33; default: 6.69)
    ha -- scaling exponent (default: 1.69)
    """
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_1_ag_p1_drainage_area_test1, points: 0.5, grade: true, solution: false -->
```python
# The visible tests below should work
from nose.tools import assert_equal

# Fake data
x = np.linspace(0.0, 20.0, 51)
area1 = drainage_area(x)

# Print results
print ('Drainage area at index 3 should be 9.1042. My drainage area is {0:.4f}.'.format(area1[3]))

# Check that the area at index 0 is half that at index 1
assert_equal(round(area1[0], 4), round(area1[1]/2.0, 4))

# Check some random values
assert_equal(round(area1[5], 4), 21.5857)
assert_equal(round(area1[-1], 4), 1057.2235)
```

<!-- grade_id: problem_1_ag_p1_drainage_area_test2, points: 0.5, grade: true, solution: false -->
```python
# The visible tests below should work

# Fake data
area2 = drainage_area(x, ka=5.0, ha=2.0)

# Print results
print ('Drainage area at index 3 should now be 7.2000. My drainage area is {0:.4f}.'.format(area2[3]))

# Check that the area at index 0 is half that at index 1
assert_equal(round(area2[0], 4), round(area2[1]/2.0, 4))

# Check some random values
assert_equal(round(area2[25], 4), 500.0)
assert_equal(round(area2[45], 4), 1620.0)
```

## Part 2: Defining a stream-power erosion function (*1 point*)

The next task in this problem is to define a function that can be used to calculate the stream-power erosion rate along the length of the river profile.
As a reminder, the stream-power erosion law has the form

\begin{equation}
  \Large
  \frac{dh}{dt} = K A^{m} S^{n}
\end{equation}

where $h$ is the elevation of the river channel, $t$ is time, $K$ is an erosional efficiency factor (accounts for lithology, climate, channel geometry, sediment supply, etc. (!)), $A$ is upstream drainage area, $S$ is channel slope, and $m$ and $n$ are the area and slope exponents.

- Create a function `calculate_dhdt` that calculates the rate of erosion along the length of the river profile using the stream-power erosion law.
    - **Note**: The value you calculate for $dh/dt$ should be negative as we'll apply it to lower the surface elevations.
- To calculate the stream-power erosion rate, you will first need to calculate the channel slope along the river profile. Take a shot at this, and have a look at [the hints for this week's exercise](https://introqg.github.io/site/lessons/L4/exercise-4.html) if you are having trouble.
- Once you have the slope, you can calculate the erosion rate $dh/dt$ using the equation above and return that value from your function

<!-- grade_id: problem_1_ag_p2_dhdt, points: 0, grade: false, solution: true -->
```python
def calculate_dhdt(elevation, dx, area, K=1.6E-8, m=1.0, n=2.0):
    """Calculate estimated drainage area upstream.
    
    Keyword arguments:
    elevation -- array of elevations along the river profile
    dx -- spacing between points along the river profile
    area -- upstream drainage area
    K -- fluvial erosional efficiency factor (units: m**0.33 / a; default: 1.6E-8)
    m -- stream power area exponent (default: 1.0)
    n -- stream power slope exponent (default: 2.0)
    """
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_1_ag_p2_dhdt_test1, points: 0.5, grade: true, solution: false -->
```python
# The visible tests below should work
from nose.tools import assert_equal

# Fake data
x1 = np.linspace(0.0, 10000.0, 101)
dx1 = x1[1]-x1[0]
topo1 = np.linspace(100.0, 0.0, len(x1))
area1 = drainage_area(x1)
dhdt1 = calculate_dhdt(topo1, dx1, area1)

# Print results
print ('dhdt at index 75 should be -0.000038. My dhdt is {0:.6f}.'.format(dhdt1[75]))

# Check max dhdt is 0.0
assert_equal(round(dhdt1.max(), 2), 0.00)

# Check some random values
assert_equal(round(dhdt1[50], 6), -1.9e-05)
assert_equal(round(dhdt1[-1], 6), -6.2e-05)
```

<!-- grade_id: problem_1_ag_p2_dhdt_test2, points: 0.5, grade: true, solution: false -->
```python
# The visible tests below should work

# Fake data
dhdt2 = calculate_dhdt(topo1, dx1, area1, K=1.E-8, m=1.0, n=1.0)

# Print results
print ('dhdt at index 75 should now be -0.002367. My dhdt is {0:.6f}.'.format(dhdt2[75]))

# Check max dhdt is 0.0
assert_equal(round(dhdt2.max(), 2), 0.00)

# Check some random values
assert_equal(round(dhdt2[25], 6), -0.00037)
assert_equal(round(dhdt2[85], 6), -0.002925)
```

## Part 3: Putting the pieces together (*0 points*)

At this point we have one of the main components for calculating the evolution of a river profile subject to stream-power erosion, the stream-power erosion function `calculate_dhdt`.
While you could go through the process of creating a number of other functions and code to handle the rest of this problem, we'll instead provide that code below.
Unfortunately, this is a slightly more involved programming exercise, so we'll go through a few other functions so that you get a sense of how the code works before you start using it.

**We strongly advise you to take the time to look through this code, rather than just running the Python cells below**. You'll be asked to modify some of this code for the other parts of the problem below.

### Defining the profile length

The function below (`create_x`) is used to define the initial length of the river profile. Review the function and check to see that you understand what it does.

- Review the code in the `create_x` function below to try to understand how it works
- Run the Python cell below

```python
def create_x(profile_length, dx):
    """Defines the initial river profile length.
    
    Keyword arguments:
    profile_length -- the length of the river profile
    dx -- the spacing between points along the river profile (units: m)
    """
    x_points = int(profile_length/dx) + 1               # Number of points for x array
    return np.linspace(0.0, profile_length, x_points)   # Return array from 0 to L by dx
```

### Defining the initial profile elevations

The function below (`create_profile`) is used to define the initial elevations along the river profile. Review the function and see whether you can understand how the 3 different topography shapes are defined, and how they differ.

- Review the code in the `create_profile` function below to try to understand how it works
- Run the Python cell below

```python
def create_profile(topo_option, x, min_elevation=0.0, max_elevation=1500.0):
    """Define the initial river profile elevations.
    
    Keyword arguments:
    topo_option -- the option for the initial river profile geometry (1-3)
    x -- array of distance along the river profile (units: m)
    min_elevation -- minimum profile elevation (units: m; default: 0.0)
    max_elevation -- maximum profile elevation (units: m; default: 0.0)
    """
    if topo_option == 1:               # "Flat" profile at 1 m elevation
        elevation = np.ones(x.size)
        elevation[-1] = 0.0
    elif topo_option == 2:             # "Plateau" profile at max_elevation
        elevation = np.zeros(x.size) + max_elevation
        elevation[-1] = min_elevation
    elif topo_option == 3:             # Sloping profile
        elevation = np.linspace(max_elevation, min_elevation, x.size)
    else:                              # Stop with error message if given a bad value
        print("Bad value listed for topo_option. It must be '1', '2', or '3'. Exiting...")
        exit()
    return elevation
```

### Putting it all together

The function below (`stream_power`) combines the functions above to calculate a river profile and plot the resulting geometry as it evolves over time. In addition to using the functions above, it also defines some of the arrays used in the calculations and creates the plots. This is slightly more complicated than the earlier function examples, but we think you should be able to read over the code and comments to get a sense of how all of the code pieces fit together.

- Review the code in the `stream_power` function below to try to understand how it works
- Run the Python cell below

```python
def stream_power(min_elevation = 0.0, max_elevation = 1500.0, profile_length = 100000.0,
                 dx = 1000.0, topo_option = 1, simulation_time = 2000000.0, dt = 2.0,
                 nplots = 5, K = 1.6E-8, m = 1.0, n = 2.0, ka = 6.69, h = 1.69):
    """Calculate evolution of a river profile with stream-power erosion.

    Keyword arguments:
    min_elevation -- minimum elevation for initial topography (units: m; default: 0.0)
    max_elevation -- maximum elevation for initial topography (units: m; default: 1500.0)
    profile_length -- length of river profile (units: m; default: 100000.0)
    dx -- spacing of points along river profile (units: m; default: 1000.0)
    topo_option -- initial topography: 1 = flat, 2 = plateau, 3 = constant slope
                   (default: 1)
    simulation_time -- total simulation time (units: a; default: 2000000.0)
    dt -- time step size for profile calculations (units: a; default: 2.0)
    nplots -- number of plots to display (default: 5)
    K -- fluvial erosional efficiency factor (units: m**0.33 / a; default: 1.6E-8)
    m -- stream power area exponent (default: 1.0)
    n -- stream power slope exponent (default: 2.0)
    h -- area scaling exponent (default: 1.69)
    ka -- distance/area scaling coefficient (units: m**0.33; default: 6.69)
    """

    # Define arrays needed for river profile calculations
    x = create_x(profile_length, dx)                        # x points for profile calcuations
    dhdt = np.zeros(len(x))                                 # erosion rate array with same size as x
    elev = create_profile(topo_option, x,                   # channel profile elevation
                          min_elevation, max_elevation)
    
    # Calculate drainage basin area as a function of distance from divide x
    area = drainage_area(x)
    
    # Define rock uplift rate and pattern
    uplift_rate = 0.001                             # Rock uplift rate [m/a]
    uplift = np.zeros(len(x)) + uplift_rate         # Create rock uplift array with constant uplift rate
    
    # Create dataframe(s) we need
    profile = pd.DataFrame(index=x, data={'Initial profile':elev})

    # Calculate plot display time interval
    dt_plot = simulation_time / nplots
    
    # Loop over simulation time and calculate river profile elevations
    time = 0.0                                           # Initialize time
    for step in range(int(simulation_time/dt)):          # Loop over amount of time between plots
        dhdt = calculate_dhdt(elev, dx, area, K, m, n)   # Calculate updated erosion rates
        outlet_elev = elev[-1]                           # Store the elevation of the river outlet
        elev += (dhdt + uplift) * dt                     # Update elevation based on uplift and river erosion
        elev[elev < 0.0] = 0.0                           # Set any elevation values below zero to 0.0
        elev[-1] = outlet_elev                           # Reset the river outlet elevation
        time += dt                                       # Increment time

        # If a profile should be plotted at this time, make a new column in the
        # DataFrame(s) to store the profile data at this time
        #
        if time % dt_plot == 0:             # Check whether the current time is an increment of dt_plot
            string=str(time)+' years'       # Create a variable with the current time
            profile[string] = elev          # Create a new DataFrame column with the current time name
    
    # Plot the results using pandas_bokeh
    ax = profile.plot_bokeh(kind='line',
                            xlim=(0, profile_length),
                            xlabel='Distance from drainage divide (m)',
                            ylabel='River channel elevation (m)',
                            title='River channel profile evolution model',
                            legend = 'bottom_left',
                            xticks=np.linspace(0.0, profile_length, 6),
                            figsize=(800,600))
```

## Part 4: Questions for this problem (*1 point*)

1. Were you able to understand how the different functions used in this problem work?
Which ones were the most challenging to understand?
Why?
2. Does it make sense why functions are used in longer programs like this one?
Do you find dividing the code into functions a useful thing to do?
If so, why?
If not, why not?

<!-- grade_id: problem_1_mg_p4_questions, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

# Problem 2: River profile evolution (*17 points*)

Now that you have a sense of how the code should work and all of the necessary functions, we can begin exploring stream-power erosion. The `stream_power` function simulates river incision into a 100-km-wide landscape with an initial flat surface elevation of 1 m. River incision is calculated using the stream-power erosion equations described in the [lecture slides for this week](https://introqg.github.io/site/_static/slides/L4/Advection.pdf).

**Notice**: Closely follow the instructions! For example, you should be sure to use **exactly** the same variable names mentioned in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.

**Your score on this problem will be based on following criteria**:

- Creating a Python function for calculating drainage area
- Creating a Python function for the stream-power erosion equation
- Answering a few questions about these functions and others defined in this problem
- Including comments that explain what most lines in the code do
- Uploading your notebook to your GitHub repository for this week's exercise

## Part 1: River erosion, visualized (*3.5 points*)

You first task in this problem is simply to run the `stream_power` function and have a look at its output. The program could take about 1 minute to run. When looking at the results, you may find it helpful to click on some of the river profile labels in the legend in order to show/hide them.

- Run the `stream_power` function in the cell below with no function parameters given
- Look at the resulting plot of the river channel profile
- Add a descriptive figure caption as if it was in a scientific journal in the cell below the plot
- Answer a few questions about the river profile model

<!-- grade_id: problem_2_mg_p1_profile, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_mg_p1_profile_caption, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

### Questions for Problem 2, Part 1 (*2 points*)

Look again through the Python code from Problem 1 and the plot it produces, and answer the following questions.

1. What is the rock uplift rate in the model? Is it constant or does it vary with space in the model?
2. What is the maximum elevation of the channel profile at the end of the simulation? Is this higher or lower than the original maximum elevation? Why?
3. Does the maximum elevation continually increase with time, or does it also decrease? Does the river profile appear to reach a steady state?
4. How fast (at what velocity) does the drainage divide (highest point in the topography) migrate across the model? To calculate this value, you should run the model several times for shorter simulation times, note the position of the divide at the completion of the simulation and then calculate the velocity (distance travelled divided by time). You can change the simulation time by providing the `simulation_time` parameter when you call the `stream_power` function. For example, `stream_power(simulation_time=10000.0)` would run the model for 10000 years.

<!-- grade_id: problem_2_mg_p1_questions, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Part 2: Subplots and erosion rates (*3.5 points*)

The code in this exercise calculates erosion rates across the length of the channel as a function of time in order to update the channel profile elevation at the end of each time step. Currently, the program plots only the profile elevations. For this part of the exercise, your goal is to plot both the elevations and erosion rates on separate plots. This can be done using the `pandas_bokeh.plot_grid()` function to add a second plot beneath the existing plot.
Currently, the first plot is created in the Python script using the commands

```python
ax = profile.plot_bokeh(kind='line',
                        xlim=(0, profile_length),
                        xlabel='Distance from drainage divide (m)',
                        ylabel='River channel elevation (m)',
                        title='River channel profile evolution model',
                        legend = 'bottom_left',
                        xticks=np.linspace(0.0, profile_length, 6),
                        figsize=(800,600))
```

This is slightly different than past plots since we're using [Pandas-Bokeh](https://github.com/PatrikHlobil/Pandas-Bokeh). In order to have two plots in a single figure we will need to use the `pandas_bokeh.plot_grid()` function, which allows us to have several plots in one figure. An example of the syntax for the `pandas_bokeh.plot_grid()` function is below:

```python
pandas_bokeh.plot_grid([[axes1], [axes2]], plot_width=800, plot_height=300)
```

This would produce a set of 2 plots on one figure, with the plot data from `axes1` displayed above the data from `axes2` (as in the example figure below). Note that there are several sets of brackets used. `plot_width` is the width of each plot in pixels and `plot_height` is the height of each plot.

![Subplot example](img/subplot_example_100ka.png)<br/>
*Figure 1. An example of using the `pandas_bokeh.plot_grid()` function. Your plot should look something like this.*

- Copy the `stream_power()` function from Problem 1 to the cell below
- Modify the copied function to have 2 subplots
    - First, create a second DataFrame called `erate` that can be used to store the erosion rates over time, and modify the code to ensure the erosion rate values are stored in the new DataFrame
        - **Note**: The erosion rate values should be positive and in units of **mm/a** in the new DataFrame
    - Then, add the code needed to generate a second plot below the first one that shows the erosion rate across the river profile
        - **Note**: In order to prevent the individual plots from being displayed separately (in addition to being together in one figure) you can use the `show_figure` parameter, described in the [Pandas-Bokeh documentation](https://github.com/PatrikHlobil/Pandas-Bokeh#output-options)
    - Be sure to update/add the axis labels
    - It will likely be helpful to have a look at the [hints for this week's exercise](https://introqg.github.io/site/lessons/L4/exercise-4.html) for this part.
- Run the test cell below the functions to produce a plot like the one above
- Add a descriptive figure caption as if it was in a scientific journal in the cell below the plot

<!-- grade_id: problem_2_mg_p2_stream_power, points: 2, grade: true, solution: true -->
```python
# Copy the stream_power function from Problem 1, paste it below, and edit for this problem

def stream_power(min_elevation = 0.0, max_elevation = 1500.0, profile_length = 100000.0,
                 dx = 1000.0, topo_option = 1, simulation_time = 2000000.0, dt = 2.0,
                 nplots = 5, K = 1.6E-8, m = 1.0, n = 2.0, ka = 6.69, h = 1.69):
    """Calculate evolution of a river profile with stream-power erosion.

    Keyword arguments:
    min_elevation -- minimum elevation for initial topography (units: m; default: 0.0)
    max_elevation -- maximum elevation for initial topography (units: m; default: 1500.0)
    profile_length -- length of river profile (units: m; default: 100000.0)
    dx -- spacing of points along river profile (units: m; default: 1000.0)
    topo_option -- initial topography: 1 = flat, 2 = plateau, 3 = constant slope
                   (default: 1)
    simulation_time -- total simulation time (units: a; default: 2000000.0)
    dt -- time step size for profile calculations (units: a; default: 2.0)
    nplots -- number of plots to display (default: 5)
    K -- fluvial erosional efficiency factor (units: m**0.33 / a; default: 1.6E-8)
    m -- stream power area exponent (default: 1.0)
    n -- stream power slope exponent (default: 2.0)
    h -- area scaling exponent (default: 1.69)
    ka -- distance/area scaling coefficient (units: m**0.33; default: 6.69)
    """
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_mg_p2_subplot, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_mg_p2_subplot_caption, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

## Part 3: Exploring the initial topography (*6 points*)

Your next task in this week's exercise is to explore the effect of the initial topography on the evolution of the river profile and erosion rates. This will hopefully give you a better sense of how the stream-power erosion equation works.

- Run the `stream_power` equation for each of the 3 initial topography options using the Python cells below
    - Run each model for 2 million years
- Below each Python cell, add a descriptive figure caption as if it was in a scientific journal in the cell below the plot
- Answer the questions below the plots

<!-- grade_id: problem_2_mg_p3_topo1, points: 0.5, grade: true, solution: true -->
```python
# Initial topography option 1 goes here

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_mg_p3_topo1_caption, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: problem_2_mg_p3_topo2, points: 0.5, grade: true, solution: true -->
```python
# Initial topography option 2 goes here

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_mg_p3_topo2_caption, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: problem_2_mg_p3_topo3, points: 0.5, grade: true, solution: true -->
```python
# Initial topography option 3 goes here

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_mg_p3_topo3_caption, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

### Questions for Problem 2, Part 3 (3 points):

Comparing the three different inintial topography options, please answer the following questions:

1. What is the starting shape of each of the initial topography options? How do they differ from one another?
2. What are the fastest erosion rates you see in any of the profiles? Which model produces the fastest erosion rates? Why do you think that might be?
3. Do the fastest erosion rates always occur in the same place, or does the location of fastest erosion change? Explain why this occurs, based on the equations for stream-power erosion presented in this week's lesson.
4. How fast does the drainage divide migrate across the model in each case? Why do you think the rates differ?
5. Look at the erosion rates across the profiles. Are the majority of the erosion rates greater than, less than or equal to the rock uplift rate? Based on this answer, is the topography in a steady state (i.e., not changing)?

<!-- grade_id: problem_2_mg_p3_questions2, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

## Part 4 - Non-uniform channel uplift (*4 points*)

Currently, the program uses a constant rock uplift rate across the river profile. Now we will explore the effects of differences in uplift velocity along the channel profile. This is equivalent to placing an active fault along the river profile. 

- Copy the `stream_power` function from above in Problem 2
- Modify it to have a rock uplift rate of 5 mm/a when x ≤ 50 km, and an uplift rate of 1 mm/a when x > 50 km
- Run your modified `stream_power` function in the Python cell below with the default parameter values
- Add a descriptive figure caption of the plot the function produces as if it was in a scientific journal in the cell below the plot
- Answer the questions below the plots

<!-- grade_id: problem_2_mg_p4_code, points: 2, grade: true, solution: true -->
```python
# Copy the stream_power function from earlier in Problem 2, paste it below, and edit for this problem

def stream_power(min_elevation = 0.0, max_elevation = 1500.0, profile_length = 100000.0,
                 dx = 1000.0, topo_option = 1, simulation_time = 2000000.0, dt = 2.0,
                 nplots = 5, K = 1.6E-8, m = 1.0, n = 2.0, ka = 6.69, h = 1.69):
    """Calculate evolution of a river profile with stream-power erosion.

    Keyword arguments:
    min_elevation -- minimum elevation for initial topography (units: m; default: 0.0)
    max_elevation -- maximum elevation for initial topography (units: m; default: 1500.0)
    profile_length -- length of river profile (units: m; default: 100000.0)
    dx -- spacing of points along river profile (units: m; default: 1000.0)
    topo_option -- initial topography: 1 = flat, 2 = plateau, 3 = constant slope
                   (default: 1)
    simulation_time -- total simulation time (units: a; default: 2000000.0)
    dt -- time step size for profile calculations (units: a; default: 2.0)
    nplots -- number of plots to display (default: 5)
    K -- fluvial erosional efficiency factor (units: m**0.33 / a; default: 1.6E-8)
    m -- stream power area exponent (default: 1.0)
    n -- stream power slope exponent (default: 2.0)
    h -- area scaling exponent (default: 1.69)
    ka -- distance/area scaling coefficient (units: m**0.33; default: 6.69)
    """
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_mg_p4_uplift, points: 0.5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6c99c9dproblem_2_mg_p4_uplift_caption, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

### Questions for Problem 2, Part 4 (*1 point*)

Looking at your plot above, please answer the following questions:

1. Is the maximum topography higher or lower in this simulation compared to that in earlier parts of this problem?
2. Does the river profile after 2 million years clearly show where the change in uplift rate occurs? Why might this be an issue in natural rivers?

<!-- grade_id: problem_2_mg_p4_questions, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE